---
title: "Einführung in R Studio"
author: CorrelAid e.V.
date: "`r Sys.Date()`"
authors:
  - Nina Hauser
  - Zoé Wolter
  - Jonas Lorenz
  - Susan Reichelt
output: 
  html_document:
    toc: true
    toc_depth: 2
    toc_float: true
    theme: flatly
    css: www/style.css
    includes:
      after_body: ./www/favicon.html
    language: de
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
knitr::opts_chunk$set(echo = FALSE, message = FALSE, warning = FALSE)
library(learnr)
library(gradethis)
source("R/setup/gradethis-setup.R")
source("R/setup/tutorial-setup.R")
# Read app parameters
params <- yaml::yaml.load_file("www/app_parameters.yml")

library(rio)
```

```{r results='asis'}
cat(get_license_note(rmarkdown::metadata$title, rmarkdown::metadata$authors))
```

Herzlich Willkommen zur dritten inhaltlichen Session unseres Kurses. Diese Woche geht es ans: Wir beginnen mithilfe von R zu programmieren. Um die R-Dateien, die wir nutzen werden, lokal zu öffnen, benötigt Ihr [R](https://ftp.fau.de/cran/){target="_blank"} und [RStudio](https://rstudio.com/products/rstudio/download/#download){target="_blank"}. Versucht gerne selbst beide Programme für Euer Betriebssystem zu installieren – idealerweise auf einem Laptop oder PC, auf dem Ihr **Adminrechte** habt. Gerne auch auf eurem privaten PC. Bei Installationsschwierigkeiten könnt Ihr uns gerne vor unserer Live-Session am Freitag **per Slack-Nachricht** kontaktieren - dann lösen wir die Herausforderungen mit Euch gemeinsam.

![*Video: Einführung in RStudio (30min)*](https://youtu.be/sgczH-e0Cxg)

# **Was ist eigentlich R und RStudio?**

## R

-   R ist eine **freie, statistische Programmiersprache**
-   **Groß- und Kleinschreibung** muss in R unbedingt beachtet werden
-   Nicht jede **Funktionalität** müssen wir in R selbst schreiben. Stattdessen nutzen wir über **`install.packages()`** und **`library()`** die Funktionalitäten, die uns andere Coder:innen frei zur Verfügung stellen - das spart viel Zeit!

## RStudio

-   RStudio ist auf R spezialisierte **integrierte Entwicklungsumgebung** (eng. Integrated Development Environment oder IDE).
-   In RStudion können wir Codes schreiben und ausführen, auf Dokumentationen und Dateien zugreifen, unsere Arbeit mit Anderen teilen uvm.
-   Wir empfehlen, den Code im **Code Editor** (zu dt. Codebearbeitungsprogramm) zu schreiben und mit "\#" zu kommentieren. Mit **"CMD + ENTER"** für Mac und **"CTRL + ENTER"** für Windows können wir den Code ausführen.
-   Die Ergebnisse unseres Codes erhalten wir in der **Console**.
-   In der **Environment** (zu dt. Umgebung) sind Objekte (also auch Daten) unserer Arbeit gespeichert.
-   Unter dem Environment gibt es einen **File Explorer** (zu dt. Dateimanager), mit dessen Hilfe wir Dateien auf unserer Festplatte suchen und per Doppelklick öffnen können.
-   Dort werden auch Plots (zu dt. Abbildungen), geladene Packages und die **Dokumentation** ("Hilfe") angezeigt, die wir mit dem Befehl **help("Funktion")** oder **?Funktion** durchsuchen können.
-   Alle **Shortcuts** für RStudio gibt es unter ("Menü Help -\> Keyboard Shortcuts Help).
-   In englischer Sprache ist zudem ein **Schummelblatt** zur [IDE RStudio](https://github.com/CorrelAid/lernplattform/blob/main/cheatsheets/03_cheatsheet-rstudio.pdf){target="_blank"} verfügbar. Diese 'Schummelblätter' gibt es zu allen möglichen Themen rund um R und RStudio, sie werden uns also im Laufe des Kurses immer wieder begleiten.

## R Projekt (RProj)

-   Wenn wir mit einer neuen Arbeit beginnen, erstellen wir am besten immer ein **R Projekt**. Das hilft uns dabei, den Überblick zu behalten und nervige Dateipfad-Verwirrungen zu vermeiden, da es als eine Art "Container" für Eure Projekte funktioniert.
-   In dem entsprechenden Ordner mit dem R Projekt legen wir deshalb am besten immer einen Ordner für unser **Daten** und einen Ordner für unseren **Code** an - falls nötig, bietet sich ein weiterer Ordner für **Grafiken** o.Ä. an.

## **Quiz**

```{r quiz_rstudio}
quiz(caption = NULL,
  question("Ich muss nicht auf die Groß- und Kleinschreibung achten, wenn ich in R programmiere.",
    answer("Korrekt"),
    answer("Inkorrekt", correct = TRUE),
    correct = "Richtig. Die Groß- und Kleinschreibung müssen wir in R stets beachten, ansonsten kann es sein, dass unser Code nicht funktionieren wird!",
    incorrect = "Leider falsch. Groß- und Kleinschreibung, aber auch Kommas und Klammern, sind wichtig damit unser Code in R funktioniert!",
    allow_retry = TRUE,
    try_again_button = "Nochmal versuchen"
  ),
  
  question("Code ist über Patent- und Urheberrechte geschützt, deshalb darf ich nichts von anderen Coder:innen benutzen.",
    answer("Korrekt, Code von anderen sollte ich auf keinen Fall verwenden. Die Nutzung anderer Codes ist immer mit rechtliche Konsequenzen verbunden."),
    answer("Inkorrekt, Code von anderen kann ich immer übernehmen - was soll schon passieren?"),
    answer("Teils teils, da R eine freie Programmiersprache ist, mit der viele Personen arbeiten, die ihre Codes öffentlich teilen, darf ich auf diese Materialien zugreifen. Allerdings sollte ich bei der Verwendung von Codes immer darauf achten, welcher Lizenz der Code unterliegt.", correct = TRUE),
    correct = "Richtig! R ist eine freie Programmiersprache und viele Menschen teilen Code zu ganz unterschiedlichen Projekten, diesen darf man grundsätzlich auch verwenden. Aber: Dabei müssen wir natürlich stets einen Blick auf die Lizenzierung des entsprechenden Codes achten und dementsprechend den Code behandeln.",
    incorrect = "Leider falsch. R ist grundsätzlich eine offene Programmiersprache mit einer großen Community, die über Foren (z.B. StackOverflow, etc.) in ständigem Austausch miteinander ist und dort auch Code zugänglich macht - rechtliche Konsequenzen müssen wir also nicht zwangsläufig befürchten. Vielmehr sind wir daran interessiert, diese wichtigen Austausch zu fördern, indem wir Code verwenden dürfen und unsere eigenen Projekte mit anderen teilen, aber Vorsicht: Wir müssen dabei immer die Lizenz eines Codes beachten - diese kann die Nutzung einschränken oder untersagen (Beispiele für Lizenzen findest Du [hier](https://creativecommons.org/about/cclicenses/)).",
    allow_retry = TRUE,
    try_again_button = "Nochmal versuchen"
  ),
  
  question("In welchem Bereich unserer 'Entwicklungsoberfläche' schreiben wir den Code?",
    answer("Konsole", correct = TRUE),
    answer("Environment"),
    answer("Terminal"),
    answer("Code Editor", correct = TRUE),
    correct = "Richtig, zum Programmieren in R verwenden wir entweder die Konsole oder den Code Editor!",
    incorrect = "Leider falsch, öffne am besten direkt nochmal RStudio. Wo findest Du die oben genannten Fenster und kannst Du dort Code eingeben?",
    allow_retry = TRUE,
    try_again_button = "Nochmal versuchen"
  ),
  
  question("IDE steht für Integrated Development Environment (zu dt. integrierte Entwicklungsumgebung). RStudio ist eine integrierte Entwicklungsumgebung für R.",
    answer("Korrekt", correct = TRUE),
    answer("Inkorrekt"),
    correct = "Richtig, R ist eine freie statistische Programmiersprache und RStudio eine auf R spezialisierte integrierte Entwicklungsumgebung (engl. Integrated Development Environment, kurz: IDE)!",
    incorrect = "Leider falsch, denn R ist einerseits eine freie statistische Programmiersprache und RStudio eine auf R spezialisierte integrierte Entwicklungsumgebung (engl. Integrated Development Environment, kurz: IDE). In diesem Fall sind also beide Aussagen richtig. Schaue Dir doch gerne weiter oben noch einmal unsere Zusammenfassung zu R, RStudio und dem R Projekt an, dort findest Du alle hilfreiche Informationen für einen erfolgreichen Start mit R!",
    allow_retry = TRUE,
    try_again_button = "Nochmal versuchen"
  )
)
```

<br>
<div style="border-width:1px; border-style:solid; border-color:#acc940; padding: 1em;">

*Jetzt geht es so richtig los! Bitte denke daran: Gerade am Anfang passieren immer viele Fehler und Du wirst vermutlich auch nicht alles auf Anhieb verstehen - das ist völlig normal und geht selbst erfahrenen Coder:innen noch so. Lass Dich davon nicht abschrecken, schreibe Dir einfach Deine Fragen auf und bringe sie zur nächsten Livesession am Freitag mit!*
</div>
<br>

# **Grundlagen in R**

## **Syntax in R**

In diesem Abschnitt stellen wir Euch die wichtigsten Grundlagen für die kommenden Wochen und Monate vor! Es gibt nämlich eine ganze Reihe mathematischer Berechnungen, die wir bereits jetzt durchführen können, ohne uns groß Gedanken über komplizierte Funktionen gemacht zu haben. Einen Überblick über die Grundlagen in R erhaltet Ihr außerdem in diesem [**Schummelblatt**](https://github.com/CorrelAid/lernplattform/blob/main/cheatsheets/03_cheatsheet-base-r.pdf){target="_blank"}. 

Starten wir also ganz simpel und verwenden R einfach als Taschenrechner! Führt mal folgenden Code-Block aus:

```{r mathop, exercise = TRUE}
## Mathematische Operatoren in R
2 + 2   # plus
2 - 1   # minus
2 * 2   # mal
2 / 2   # geteilt
```

Wir können R also als Taschenrechner verwenden und dabei von den Grundregeln bis hin zu komplizierten Formeln alles mögliche anstellen. Ihr seht außerdem, dass wir **Kommentare** in unseren Code integrieren können - Ihr erkennt sie an dem Hash **"\#"**. Das bedeutet, dass alles was in der Zeile nach einem **"\#"** kommt, von R nicht als Code ausgeführt wird!

Mit **logischen Ausdrücken** können wir allerdings noch viel mehr machen, als solche Berechnungen. Wir können R nämlich auch als eine Art Orakel verwenden. Mithilfe der Ausdrücke können wir u.a. Datensätze filtern und Wenn-Dann-Begingungen für kompliziertere Operationen definieren. Aber starten wir einmal einfach: Was denkt Ihr, sind die folgenden logischen Ausdrücke wahr (=TRUE) oder falsch (=FALSE)?

```{r logik, exercise = TRUE}
## Logische Ausdrücke in R
2 == 2                      # ist gleich
2 != 3                      # ist ungleich
!(2 == 2)                   # nicht
2 < 3                       # kleiner als
2 > 3                       # größer als
2 <= 2                      # kleiner gleich
2 >= 3                      # größer gleich
(1 + 1 == 2) & (2 + 2 == 5) # und
(1 + 1 == 2) | (2 + 2 == 5) # oder
```

## **Datentypen in R**

Was für Daten gibt es denn eigentlich? Erinnern wir uns dazu einmal an letzte Woche zurück, als wir uns verschiedene Datentypen angeschaut und überlegt haben, welche (un-)geordneten Informationen in Daten überhaupt enthalten sein können. Diese **Unterscheidung der Datentypen** werden an dieser Stelle wieder relevant, denn in R entscheidet auch die Art der Daten darüber, was wir überhaupt mit ihnen anstellen können. An dieser Stelle möchten wir daher einmal einen Blick auf die Grundtypen an Daten werfen, die es in R gibt und mit denen wir in Zukunft arbeiten werden:

```{r grundtypen, exercise = TRUE}
## Grundtypen an Daten in R
7                # integer
7.234123         # numeric
'irgendein Text' # character
TRUE             # logical
```

  - **numeric**: Alle Zahlen, die einen Dezimalwert haben oder einen Bruch darstellen.
  - **integer**: Ganze Zahlen, die keinen Dezimalwert haben.
  - **character**: Jegliche Form der Text-Daten. Diese müssen immer in Anführungszeichen stehen.
  - **logical**: Eine Variable mit binärer Ausprägung (auch: boolean), also z.B. das binäre  TRUE oder FALSE.

## **R als Gedächtnis**

Man kann durchaus sagen, dass R über eine Art Gedächtnis verfügt. Denn mithilfe von R können wir **Werte in Objekten speichern**. In den entsprechenden Operationen verwenden wir dafür den Zuweisungspfeil `<-`, mit dem wir ganz simpel einem Objekt eine Zahl zuweisen. In diesem Fall weisen wir dem Objekt 'zahl' den Wert $2$ zu. Wenn wir anschließend das Objekt wieder aufrufen, wird uns der Wert angezeigt, den wir zuvor im Objekt gespeichert haben.  

```{r zuweisung, exercise = TRUE}
## Variablenzuweisung zu einem Objekt
zahl <- 2       # Variablenzuweisung
zahl            # Aufrufen des Objekts
```

Nun denkt ihr vielleicht, dass ein Objekt, bestehend aus nur einem Wert, ja ziemlich übersichtlich ist und man dafür vielleicht gar kein R braucht ..., aber neben einzelnen Werten lässt sich natürlich auch eine **ganze Reihe an Werten in einem Objekt speichern**. Hierbei handelt es sich um einen sogenannten **Vektor**. Mithilfe des alle `c()` können wir dabei alle einzelnen Werte zusammengepacken, aber Achtung: In einem Vektor sollten nur Werte des **gleichen Grundtyps** gespeichert sein! 
In diesem Fall werden wir also zuerst alle einzelnen Werte des selben Typs in einem Vektor mit der Bezeichnung 'irgendein_vektor' speichern und anschließend aufrufen. Wie wir sehen, werden uns anschließend alle vier Werte, die wir in dem Vektor gespeichert haben, angezeigt.
Glücklicherweise können wir auch einzelne Werte, die in einem Vektor enthalten sind, wieder aufspüren. Dazu nutzen wir die Position des Wertes im Vektor, die wir mithilfe der eckigen Klammern '[]' angeben. 

```{r vector, exercise = TRUE}
## Erstellen eines Vektors
irgendein_vektor <- c('R', 'ist', 'richtig', 'toll!') # Speichern der Werte im Vektor  
irgendein_vektor                                      # Aufrufen des Vektors

## Aufspüren eines bestimmten Wertes innerhalb eines Vektors
irgendein_vektor[1] 
irgendein_vektor[4]
```

Einen letzten Datentyp neben den vier Grundtypen und Vektoren möchten wir Euch hier noch vorstellen: Listen - in Listen können Werte von unterschiedlichen Grundtypen gespeichert werden:

```{r list, exercise = TRUE}
# Einzelne Werte werden in einer Liste mit der Bezeichnung 'meine_liste' gespeichert, wenn Ihr diese aufruft, werden alle darin enthaltenen Werte angezeigt
meine_liste <- list('Das', 'ist', 'eine', 'Liste', 1, 5.23, TRUE)
meine_liste

# Natürlich könnt Ihr auch einen einzelnen Wert, der in einer Liste enthalten ist, wieder aufspüren, dazu nutzt Ihr einfach die Position des Wertes in der Liste
meine_liste[1]
```

In der Environment behaltet ihr einen ganz guten Überblick über die verschiedenen Daten und Objekte, mit denen ihr in R arbeitet. Aber ihr könnt auch Funktionen nutzen, um euch einen schnellen Überblick zu verschaffen. Wie genau die Objekte im Detail aufgebaut sind, schauen wir uns gleich noch an!

```{r class, exercise = TRUE}
# Um herauszufinden mit was für einem Objekt ihr umgeht, könnt ihr die class() Funktion anwenden.
beispiel1 <- 5+4
beispiel2 <- c('Hund', 'Katze')
beispiel3 <- list('Auch', 'das', 'ist', 'eine', 'Liste', TRUE)

class(beispiel1)
class(beispiel2)
class(beispiel3)
```

### **Interaktive Übung**

<div style="border-width:1px; border-style:solid; border-color:#acc940; padding: 1em;">
*Hinweis: In den folgenden Übungslektionen gibt es einen Teil, den Ihr innerhalb der Lernplattform erledigen könnt. Den anderen Teil solltet Ihr lokal auf Eurem Rechner ausführen, um Euch an die IDE RStudio zu gewöhnen.*
</div>
<br>

Wenn Ihr mit Eurem Datenprojekt startet, solltet Ihr Euch überlegen, welche **Funktionalitäten** Ihr braucht, um Daten zu importieren, zu bereinigen und zu analysieren. Die benötigten Funktionalitäten müsst Ihr nicht alle selbst programmieren, dafür gibt es sogenannte **Packages** (zu dt. Pakete) - eine ziemlich gute Sache. Eure Arbeit beginnt deshalb in der Regel mit der **Installation und dem Laden** dieser Funktionalitäten. Dazu nutzen wir die Funktionen **`install.packages("package")`** und **`library(package)`**. Die Nutzung von Anführungszeichen ist nur bei der Installation (also für `install.packages()`) relevant. Bitte denkt daran, `install.packages()` auszukommentieren (also ein \# vor die Zeile zu packen), wenn die Installation abgeschlossen ist - sonst wird die Installation bei jedem Ausführen des Skripts wieder durchgeführt. In der Zukunft könnt Ihr auch eine eigene Datei für Installationsskripte einrichten - damit könnt Ihr Euch aber noch zu einem späteren Zeitpunkt beschäftigen.

Ein **Code Chunk** (zu dt. Codeblock), also ein Abschnitt mit Code, der Pakete installiert und lädt, sieht so aus. Hier wurde der Code auskommentiert (da wir alle Pakete, die für die Lernplattform notwendig sind für Euch schon installiert und geladen haben).

```{r pakete, exercise = TRUE}
# Installation und Laden des Packages "rio"
# install.packages("rio")
library(rio)
```

Mit der Notation **`package::funktion()`** kennzeichnen wir, woher eine Funktion kommt und können - unabhängig von geladenen Package-Libraries - Funktionen nutzen. Dabei ist es zwar nicht nötig, das Package vor die Funktion zu schreiben - aber definitiv best practice! Es gibt unglaublich viele Packages, die jeweils auch viele Funktionen mitbringen. Für viele Vorgehensweisen gibt es demnach nicht "einen" richtigen Weg, sondern je nach Package und Funktionen relativ viel Variation. Das Hinzufügen aus welchem Package eine bestimmte Funktion stammt, hilft also auch dabei, den eigenen Überblick über den gewählten Ansatz zu behalten.

Dateien können u.a. mit dem Package **`rio`** und der Funktion **`import()`** geladen werden, wie ihr im Code Chunk unten sehen könnt. Die geladenen Daten werden dann in einem Dataframe gespeichert - hier wird dem Datensatz der Name "data_raw" zugewiesen ("objektname \<- objekt"). Über diesen ist er dann im weiteren Skript abrufbar. Wenn Ihr Daten aus lokalen Ordnern laden möchtet, müsst Ihr den Link durch den Pfad der Datei (einsehbar über Rechtsklick auf den Dateinamen) auf Eurem Rechner ersetzen. Leichter wird das mit der Funktion **`here::here()`**, die für Euch in Verbindung mit RProjekten die Navigation in den entsprechenden Ordner stark vereinfacht (Recap: die `here`-Funktion setzt das Arbeitsverzeichnis in den Ordner mit der .Rproj-Datei).

```{r daten_einlesen, exercise = TRUE}
# Hier laden wir mit dem rio-Package und der Funktion "import()" unseren Datensatz.
data_raw <- rio::import('https://raw.githubusercontent.com/rfordatascience/tidytuesday/master/data/2021/2021-01-26/plastics.csv')
# Wenn Ihr das Objekt aufruft, bekommt Ihr die darin gespeicherten Daten angezeigt
data_raw


# Alternative für lokale Dateien, die in Eurem R-Projekt als CSV-Datei in einem Ordner namens "Daten" hinterlegt wurden, mit dem "here"-Package
# data_raw <- rio::import(here::here('daten/plastics.csv'))
```

Überprüft hier nochmal, als welche Art von **Objekt** `rio::import()` unseren Datensatz importiert. Das haben wir eben schon mit anderen Objekten probiert ... Wenn Ihr auf "Antwort einreichen" klickt, bekommt Ihr Feedback zu Eurem Code. Mit Hilfe von "Lösung anzeigen" könnt Ihr nachschauen, wie die richtige Lösung aussieht.

```{r exercise_class, exercise = TRUE}
# Euer Code hier
```

```{r exercise_class-solution}
# Datenformate überprüfen wir mit "class()"
class(data_raw)
```

```{r exercise_class-check}
grade_this_code()
```

Mit der Funktion **`head()`** könnt Ihr Euch die ersten sechs Beobachtungen anzeigen lassen, um den Datensatz zu erkunden. Das Gegenteil von `head()` ist `tail()` - damit könnt Ihr die letzten sechs Beobachtungen erkunden.

```{r head, exercise = TRUE}
# Erste sechs Zeilen anzeigen lassen
head(data_raw)
```

Um eine Übersicht über die Variablen zu bekommen, könnt Ihr Euch mit der Funktion **`colnames()`** deren Namen anzeigen lassen.

```{r variablen, exercise = TRUE}
# Spaltennamen anzeigen lassen
colnames(data_raw)
```

Die Funktionen **`nrow()`, `ncol()` und `dim()`** helfen Euch die **Dimensionen** Eures Datensatzes zu bestimmen. Mit `help("Funktion")` könnt Ihr Euch über RStudio in der Konsole die Erklärung zu den Funktionen anzeigen lassen. Nutzt das bisher Gelernte, um Euch die Spaltenanzahl (also die Anzahl an Variablen) ausgeben zu lassen.

```{r exercise_ncol, exercise = TRUE}
# Euer Code hier
```

```{r exercise_ncol-solution}
# Mit ncol() könnt Ihr Euch die Anzahl an Spalten (eng. columns) anzeigen lassen
ncol(data_raw)
```

```{r exercise_ncol-check}
  give_code_feedback(grade_this({
    pass_if_equal(.solution)
    if (length(.result) != 1) {
        fail("Eure Antwort enthält {length(.result)} Wert(e) und nicht: {length(.solution)}. Habt Ihr die richtige Funktion genutzt?")
    }
    if (.result == 13380) {
        fail("Das ist die Zeilenanzahl, also die Anzahl an Beobachtungen! Habt Ihr die richtige Funktion genutzt?")
    }
    fail("Noch nicht ganz richtig.")
  }))
```

Elementare Merkmale des Datensatzes könnt Ihr Euch mit der Funktion **`summary()`** anzeigen lassen. Bei Text sind das Länge, Objektart und Mode; bei numerischen Werten Minimum, Mittelwert, Maximum, Quartile und fehlende Werte (diese Begriffe sollten euch aus der letzten Lektion bekannt vorkommen - wir werden in der nächsten Lektion auch nochmal genauer auf die Bedeutung schauen).

```{r datenstruktur_summary, exercise = TRUE}
# Überblick über Daten
summary(data_raw)
```

Alternativ gibt es auch die Funktion **`str()`**, die ähnlich wie **`summary()`** Grundeigenschaften von Datentabellen anzeigt. Welche Unterschiede findet Ihr zwischen den beiden Funktionen?

```{r datenstruktur_str, exercise = TRUE}
# Überblick über Daten
str(data_raw)
```

Wenn Ihr einzelne Spalten oder **Subsets** (zu dt. Teilmengen) eines Datensatzes braucht, könnt Ihr diese aus dem Datensatz herausziehen. Dazu nutzen wir das Dollarzeichen "\$" mit Datensatz\$Spaltenname oder die eckigen Klammern mit Datensatz['Spaltenname']. Wollt Ihr mehrere Spalten aufrufen, könnt Ihr die Logik Datensatz[c('Spaltenname1', 'Spaltenname2')] nutzen (Recap: **`c()`** ist hierbei die Notation von R für Vektoren, in denen Ihr mehrere Werte hinterlegen könnt).

Wenn Ihr mit einem bestimmten Subset weiterarbeiten wollt, könnt ihr auch einfach über eine neue Objektzuweisung und einen Objektnamen einen zusätzlichen Datensatz erstellen.

```{r slicing1, exercise = TRUE}
# Auslesen der Spalte 'country'
data_raw$country 
# Auslesen der Spalten 'country' und 'year'
data_raw[c('country', 'year')]
```

```{r slicing2, exercise = TRUE}
# Erstellung eines Subsets auf Basis von Bedingungen
subset(data_raw, country == "Argentina" & year == 2020)
```

```{r slicing3, exercise = TRUE}
# Erstellung eines neuen Objekts auf Basis von Bedingungen
ARG_2020 <- subset(data_raw, country == "Argentina" & year == 2020)
```

<div style="border-width:1px; border-style:solid; border-color:#acc940; padding: 1em;">

*Uff geschafft! Wenn Du direkt noch mehr zu R lernen möchtest, dann schau Dir gerne den Exkurs zu Kontrollstrukturen an. Gerne kannst Du aber auch direkt weiter zur Übung springen!*

</div>

---

<details>
  <summary><h4>&#10145; Exkurs: Kontrollstrukturen</h4></summary>
  <br>
  
<h3><b>Exkurs: Kontrollstrukturen</b></h3>  

In der Programmierung gibt es verschiedene Kontrollstrukturen: Nur *wenn* die Note eines Teilnehmenden besser als eine 2 ist, *dann* soll R den Text 'Gut gemacht' ausgeben? Dann braucht Ihr eine **Wenn-Dann-Bedingung** (*engl. If-Then-Statement*). Hier die Syntax und ein kleines Beispiel - Welcher Text (*Operation innerhalb der geschwungenen Klammern*) wird hier ausgegeben?

```{r wenndann, exercise = TRUE}
# if (Bedingung, die überprüft wird) {
#   Code, der ausgeführt wird, wenn die Bedingung erfüllt ist
# } else {
#   Code, der ausgeführt wird, wenn die Bedingung NICHT erfüllt ist
# }

if (2 + 2 == 5) {
  print("Oh nein!")
  } else if (2 + 2 <= 3) {
  print("Oh nein!")
  } else if (2 + 2 >= 5) {
  print("Oh nein!")
  } else {
  print("Puh, nicht verrechnet...")
    }
```

Da haben wir nochmal Glück gehabt: Unser **regelbasiertes System** funktioniert. Neben **Maschinellem Lernen** sind solche regelbasierten Systeme übrigens eine wichtige Form von allen Technologien, die unter das Modewort **Künstliche Intelligenz** fallen. Die zugehörigen Regeln werden in komplexeren Systemen von Expert:innen verfasst und erst nach und nach, wenn es möglich ist, durch alternative Technologien ersetzt. Ein gutes Beispiel dafür sind Verfahren zur Verarbeitung von Sprache (siehe [hier](https://aclanthology.org/D13-1079.pdf){target="_blank"}). Ihr könnt Euch aber auch einfach ein regelbasiertes System vorstellen, das automatisch bei der Erreichung einer bestimmten Spendensumme eine Dankesnachricht an alle Spender:innen versendet.

Neben der Ausführung von Aktionen unter bestimmten Bedingungen können wir durch technische Implementierungen **Aktionen mehrfach ausführen**. Ein Beispiel: Nach der Erreichung der Spendensumme soll eine Dankesnachricht an alle Spender:innen gesendet werden. In der Programmierung bezeichnet man diese iterierenden Verfahren als **Loop** (zu dt. Schleife). Sind alle Spender:innen bekannt (und somit die Liste der Elemente, über die man iteriert), nutzt man einen **For Loop** (zu dt. Für-Schleife).

![](https://github.com/CorrelAid/lernplattform/blob/main/abbildungen/03_einfuehrung-in-rstudio/forloop.png?raw=true){#id .class width="30%" height="100%"}

```{r fuerschleife, exercise = TRUE}
# Aus der Liste "list("Lisa", "Frie", "Zoe")" picken wir uns eine "spenderin" raus...
for (spenderin in list("Lisa", "Frie", "Zoe")) {
  # für diese "spenderin" geben wir den Satz 'Danke, <Name diese:r:s Spender:in>!' aus
  print(paste("Danke,", spenderin, "!"))
}
```

Soll die Aktion ausgeführt werden, bis eine Bedingung eintritt, z.B. weil alle Spender:innen kontaktiert wurden, dann nutzt man einen **While Loop** (zu dt. Solange-Schleife).

![](https://github.com/CorrelAid/lernplattform/blob/main/abbildungen/03_einfuehrung-in-rstudio/whileloop.png?raw=true){.class width="30%" height="100%"}

```{r solangeschleife, exercise = TRUE}
# Erstellung der Liste unbekannter Länge
spenderinnenliste <- list("Lisa", "Frie", "Zoe")

# definiere Startwert
i <- 1

# führe den Code für alle i kleiner gleich der Länge der Liste durch...
while (i <= length(spenderinnenliste)) {
  # ...gebe für diesen Eintrag den Satz 'Danke, <Name diese:r:s Spender:in>!' aus...
  print(paste("Danke,", spenderinnenliste[[i]], "!"))
  # ...setze den Counter i um 1 herauf
  i = i + 1
}
```

</details>

---

Eine letzte Anmerkung: Es ist ganz normal, dass Euer Code mal **Fehler** wirft - auch uns passiert das. In der Console (und in RMDs unter dem Code Chunk) wird Euch dann angezeigt, **wo der Fehler auftritt und um welche Art von Fehler es sich handelt** (sog. Fehlermeldung). Beliebte Fehlerquellen sind **Syntaxfehler** (wenn Ihr beispielsweise Klammern nicht schließt, die falsche Klammerart nutzt oder Argumente nicht durch Kommata trennt) oder die falsche Verwendung von Argumenten in Funktionen. Auch das Schreiben von Kommentaren in Code Chunks ist nur mit "\#" zu Beginn des Textes möglich. Wenn Ihr mal gar nicht weiter wisst, kopiert einfach die Fehlermeldung aus der Console und begebt Euch damit auf eine Googlesuche. Auf [**Stack Overflow**](https://stackoverflow.com/){target="_blank"} und in der [**R-Online Community**](https://community.rstudio.com/){target="_blank"} finden sich für (fast) alle Herausforderungen Antworten und Hilfestellungen.


### **Und jetzt Ihr**

Legt nun ein neues R Projekt (wie im Video) und eine neue R Markdown-Datei auf Eurem Rechner an und...

1.  Installiert und ladet das Package "rio".
2.  Lest den Datensatz über den Hyperlink oder lokal ein. Ihr findet ihn unter diesem [Link](https://raw.githubusercontent.com/rfordatascience/tidytuesday/master/data/2021/2021-01-26/plastics.csv){target="_blank"} und könnt ihn dort auch herunterladen, wenn Ihr Euch für die zweite Variante entscheidet.
3.  Nutzt die vorgestellten Funktionen, um den Datensatz zu erkunden. Dafür könnt Ihr die Code Chunks aus der Lernplattform einfach in RStudio kopieren und sie dort ausführen.

### **Zusätzliche Ressourcen**

-   Schummelblatt zu den [Grundlagen in R](https://github.com/CorrelAid/lernplattform/blob/main/cheatsheets/03_cheatsheet-base-r.pdf){target="_blank"} (dt.)
-   Schummelblatt zur [IDE RStudio](https://github.com/CorrelAid/lernplattform/blob/main/cheatsheets/03_cheatsheet-rstudio.pdf){target="_blank"} (engl.)
-   [Introduction to Programming in R (engl.)](https://app.dataquest.io/course/intro-to-r-rewrite){target="_blank"} von Dataquest
-   [Data Structures in R (engl.)](https://app.dataquest.io/course/datastructure-in-r-rewrite){target="_blank"} von Dataquest
-   [Control Flow, Iteration and Functions in R (engl.)](https://app.dataquest.io/course/intermediate-r){target="_blank"} von Dataquest
-   [R Studio Learn: Programming Basics](https://rstudio.cloud/learn/primers/1.2){target="_blank"}

<a class="btn btn-primary btn-back-to-main" href=`r params$links$end_session`>Session beenden</a>
